<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE09</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber>MA_LE09</idnumber>
  </question>

<!-- question: 1257772  -->
  <question type="ddwtos">
    <name>
      <text>MA_LE09_22_Stetigkeit bei Grenzwerten</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Berechnen Sie den folgenden Grenzwert:
$$\lim\limits_{x \to 0^+} \ln\left(\frac{x^2+4x}{3x^3+2x^2+2x}\right)$$<br>
<br>
Da die Logarithmusfunktion [[1]] ist, können wir den Limes an der Funktion nach Innen verschieben.
Wir erhalten:
$$\lim\limits_{x \to 0^+} \ln\left(\frac{x^2+4x}{3x^3+2x^2+2x}\right)= \ln\left(\lim\limits_{x \to 0^+}\frac{x^2+4x}{3x^3+2x^2+2x}\right)$$
Wir klammern [[3]] aus dem Nenner und Zähler des Bruchs aus. Da der Grenzwert gegen null läuft, sind dies die [[4]] Potenzen. Es gilt:

$$\frac{x^2+4x}{3x^3+2x^2+2x}= \frac{x}{x}\cdot \frac{x+4}{3x^2+2x+2}$$
Mit [[6]] folgern wir, dass
$$\lim\limits_{x \to 0^+} \frac{x+4}{3x^2+2x+2}= \frac{4}{2}=2$$
gilt. Eingesetzt in die Aufgabe erhalten wir:
$$\lim\limits_{x \to 0^+} \ln\left(\frac{x^2+4x}{3x^3+2x^2+2x}\right)= \ln\left(\lim\limits_{x \to 0^+}\frac{x^2+4x}{3x^3+2x^2+2x}\right)=\ln(2)$$]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MA_LE09_22</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[Prima, Ihre Antworten sind richtig!<br>
[stashdrop secret="iv4122" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>stetig</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>positiv</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>führende Terme</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>niedrigen</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>hohen</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>den Rechenregeln</text>
      <group>1</group>
    </dragbox>
  </question>

<!-- question: 839728  -->
  <question type="multichoice">
    <name>
      <text>MAT_LE09_11_Stetigkeit_Def_umformulieren</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Welche der nachstehenden Gleichungen bedeuten, dass die Funktion \(f: \mathbb{R} \to \mathbb{R}\) im Punkt \(x=a\) stetig ist?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_11</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text><![CDATA[Prima, Ihre Antwort ist richtig!<br>
[stashdrop secret="Zg6BlB" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(a)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(\lim\limits_{x \to a}x)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="html">
      <text><![CDATA[<p>\(f(a) = f(\lim\limits_{x \to a}x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a^+}f(x) =\lim\limits_{x \to a^-}f(x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <hint format="markdown">
      <text>**Hinweis 1:**

Schauen Sie sich noch einmal die Definition der Stetigkeit an und überprüfen, welche Unterschiede die Gleichungen zur Definition aufweisen.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 2:**

Überlegen Sie sich Beispiele für die einzelnen Gleichungen. Können Sie Funktionen konstruieren, die in $x=a$ nicht stetig sind?</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 3 - Musterlösung:**

Die Gleichungen
$$\lim\limits_{x \to a}f(x) = f(a)$$
 und
$$\lim\limits_{x \to a}f(x) = f(\lim\limits_{x \to a}x)$$
sind richtig.</text>
    </hint>
  </question>

<!-- question: 839729  -->
  <question type="multichoice">
    <name>
      <text>MAT_LE09_13_Stetig_im_Definitionsbereich</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Die Funktion \( f(x) = \frac{1}{x-5} \) ist stetig auf ihrem Definitionsbereich \(D=\mathbb{R}\setminus\{5\}\).</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_13</idnumber>
    <single>false</single>
    <shuffleanswers>false</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text><![CDATA[Prima, Ihre Antwort ist richtig!<br>
[stashdrop secret="DHKusv" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Leider ist Ihre Antwort falsch. Versuchen Sie es noch einmal!</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="100" format="html">
      <text>Wahr</text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-100" format="html">
      <text>Falsch</text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <hint format="html">
      <text><![CDATA[<strong>Hinweis: </strong>
<br>
Die Antwort ist falsch. Die Funktion \( f \) ist für alle \( x \not= 5 \) definiert und der Graph kann außerhalb dieses Punktes durchgezeichnet werden. Dies bedeutet, dass \( f: \mathbb{R} \setminus \{5\} \to \mathbb{R} \) auf dem gesamten Definitionsbereich stetig ist.]]></text>
    </hint>
  </question>

<!-- question: 1257782  -->
  <question type="multichoice">
    <name>
      <text>MAT_LE09_4_MC_Frage</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Ist die Funktion f stetig?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_4</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text><![CDATA[Ihre Antwort ist richtig. Nach dieser Lerneinheit werden Sie dies auch präzise erklären können.<br>
[stashdrop secret="FiD44q" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Ihre Antwort ist falsch. Versuchen Sie es noch einmal!</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="-30" format="html">
      <text><![CDATA[<p>Ja, weil wir die Funktion zeichnen können.</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="100" format="html">
      <text><![CDATA[<p>Nein, weil wir den Stift beim Zeichnen bei x=0 absetzen müssen.</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="markdown">
      <text><![CDATA[<p>Nein, weil wir den Stift beim Zeichnen bei x=0 und x=1 absetzen müssen.</p>]]></text>
      <feedback format="markdown">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="markdown">
      <text><![CDATA[<p>Ja, weil F für alle reellen Zahlen definiert ist.</p>]]></text>
      <feedback format="markdown">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257771  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_19_Stetige Fortsetzbarkeit</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Entscheiden Sie für die folgenden Funktionen stetig fortsetzbar sind:<br>
<table>
<tr>
<td> Die Funktion \(f: \mathbb{R} \setminus\{2\} \to \mathbb{R}\) <br><br>
<img src="/theme/boost_union_vorsprung/pix/static/mathematik/le09/mat-le09-19b-grafik-beispiel-fuer-aufgabe-1.png" width="75%">
<br><br>
ist in \(x=2\) [[input:ans1]].[[validation:ans1]][[feedback:prt1]]</td>

<td> \(\quad\)    </td>

<td>Die Funktion \(f: \mathbb{R} \setminus\{0\} \to \mathbb{R}\) <br><br>
<img src="/theme/boost_union_vorsprung/pix/static/mathematik/le09/mat-le09-19c-grafik-beispiel-fuer-aufgabe-2.png" width="75%">
<br><br>
ist in \(x=0\) [[input:ans2]].[[validation:ans2]][[feedback:prt2]]</td>

</tr>
</table></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_19</idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[L1: [[1,true, "stetig fortsetzbar"],[2,false,"nicht stetig fortsetzbar"]];
L2: [[1,false, "stetig fortsetzbar"],[2,true,"nicht stetig fortsetzbar"]];

L1:random_permutation(L1);
L2:random_permutation(L2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt3]]</text>
    </specificfeedback>
    <questionnote>
      <text>L1:random_permutation(L1);
L2:random_permutation(L2);
L3:random_permutation(L3);
L4:random_permutation(L4);
L5:random_permutation(L5);
L6:random_permutation(L6);</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>L1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>L2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Richtig, die Folge ist stetig fortsetzbar. Wir wählen in \(x=2\) als Funktionswert den Wert \(4\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Ihre Antwort ist leider falsch. </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Richtig, die Funktion ist nicht stetig fortsetzbar. Den Wert \(\infty\) können wir nicht als Funktionswert setzen.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Ihre Antwort ist leider falsch.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Prima!<br></p>
[stashdrop secret="50hSKi"  image]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 1:</strong></p>

<p>Wiederholen Sie die Definition der stetigen Fortsetzbarkeit noch einmal.</p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 2:</strong></p>

<p>Welchen Punkt würden Sie im Graphen jeweils ergänzen wollen, um den Graphen durchzeichnen zu können?</p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 3 - Musterlösung:</strong></p>

<p> Die erste Funktion ist stetig fortsetzbar. In die Lücke bei \(x=2\) ergänzt man als Funktionswert den Wert \(4\).<br>
Die zweite Funktion ist nicht stetig fortsetzbar. Den Wert \(\infty\) können wir nicht als Funktionswert setzen.</p>]]></text>
    </hint>
  </question>

<!-- question: 1257040  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_24_Parameterabhängige_Stetigkeit</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>Bestimmen Sie \(a\) und \(b\) so, dass die Funktion</p>
<p>
\(f(x)&nbsp; = \left\{ \begin{array}{ll} {@fterm@} &amp; \textrm{ für }\; x\leq {@l@} \\
 {@gterm@} &amp;\textrm{ für }\; {@l@} &lt; x\leq {@m@} \\
{@hterm@} &amp; \textrm{ für }\;  x &gt; {@m@}
\end{array}\right.\)</p>
<p>stetig ist.
  <br>
</p>
<p>Antwort:
  <br>
</p>
<p>a = [[input:ans1]]
  <br>
</p>
<p>b = [[input:ans2]]
</p>
<p>[[validation:ans1]]
  <br>
</p>
<p>[[validation:ans2]]</p>
[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_24</idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>j:2;    /* bestimmt einen der Koeffizienten von f   */
k:1;   /* bestimmt einen der Koeffizienten von h  */
l:-1;   /* entscheidet, wo der Übergang von f zu g stattfindet   */
m:1;   /* entscheidet, wo der Übergang von g zu h stattfindet   */
f(x):=x^2+j*a*x+b;
fterm:x^2+j*a*x+b;
 g(x):=(a+b)*x;
gterm:(a+b)*x;
 h(x):=x^2-a*x+k*b;
hterm:x^2-a*x+k*b;
c:flatten(solve([f(l)=g(l),g(m)=h(m)],[a,b]));
aloes:rhs(c[1]);
bloes:rhs(c[2]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt2]]</text>
    </specificfeedback>
    <questionnote>
      <text>Stetigkeit</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>aloes</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>bloes</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>aloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(a={@aloes@}\) ist richtig bestimmt.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(a\) ist nicht richtig bestimmt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(b={@bloes@}\) ist richtig bestimmt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(b\) ist nicht richtig bestimmt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(b={@bloes@}\) ist richtig bestimmt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(b\) ist nicht richtig bestimmt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>l^2+j*ans1*l+ans2-(ans1+ans2)*l</sans>
        <tans>0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@l@}\) stetig.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@l@}\) nicht stetig.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>(ans1+ans2)*m-(m^2-ans1*m+k*ans2)</sans>
        <tans>0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@m@}\) stetig.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@m@}\) nicht stetig.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>aloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[Prima!<br>
[stashdrop secret="60Diy1"  image]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <hint format="html">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Wählen Sie die Parameter \( a \) und \( b \) so, dass die Funktion \(f\) stetig ist. Dafür müssen die Funktionswerte an den Übergängen der Definitionsbereiche \(x = -1 \) und \( x = 1 \) übereinstimmen. </p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Berechnen Sie die Funktionswerte an den Übergängen der Definitionsbereiche in Abhängigkeit der Parameter \( a \) und \( b \), indem Sie \(-1\) beziehungsweise \( 1 \) in \(x\) einsetzen. Finden Sie nun geeignete \( a \) und \( b \), so dass die Werte jeweils für die ersten beiden Gleichungen für \( x = -1 \) und die zweite und dritte Gleichung für \(x = 1 \) übereinstimmen.]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung: </b> </p>
<p> An der Stelle \( x = -1 \) gilt: \( f(-1) = (-1)^2 + 2 \cdot a \cdot (-1) + b = 1 - 2 \cdot a + b \). Damit \( f \) an dieser Stelle stetig ist, muss ebenfalls gelten: \( f(-1) = (b+a) \cdot (-1) = -a-b \). Also muss gelten: \( 1 - 2 \cdot a + b = -a-b \) und somit \( 1 = a -2b \). </p>
<br>
<p>Analog gilt an der Stelle \( x = 1 \): \( f(1) = (b+a) \cdot 1 = a + b \). Damit \(f \) an dieser Stelle stetig ist, muss gelten: \( f(1) = 1^2 - a \cdot 1 + b = 1-a+b \). Also muss auch gelten: \( a + b = 1-a+b \) und somit \( 1 = 2a \). Aus der zweiten Gleichung folgt also, dass \( a = \frac{1}{2}\) ist. Setzt man diese Lösung in die erste Gleichung für \( a \), so erhält man: \( b = -\frac{1}{4} \). </p>]]></text>
    </hint>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE09/Aufgaben zur Übung</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 1257777  -->
  <question type="ddwtos">
    <name>
      <text>MA_LE09_22_Stetigkeit bei Grenzwerten (Variante)</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[Berechnen Sie den folgenden Grenzwert:
$$\lim\limits_{x \to \infty} \ln\left(\frac{x^3+4x}{3x^3+2x^2+2x}\right)$$<br>
<br>
Da die Logarithmusfunktion [[1]] ist, können wir den Limes an der Funktion nach Innen verschieben.
Wir erhalten:
$$\lim\limits_{x \to \infty} \ln\left(\frac{x^3+4x}{3x^3+2x^2+2x}\right)= \ln\left(\lim\limits_{x \to \infty}\frac{x^3+4x}{3x^3+2x^2+2x}\right)$$
Wir klammern [[3]] aus dem Nenner und Zähler des Bruchs aus. Da der Grenzwert gegen unendlich läuft, sind dies die [[5]] Potenzen. Es gilt:

$$\frac{x^3+4x}{3x^3+2x^2+2x}= \frac{x^3}{x^3}\cdot \frac{1 + \frac{4}{x^2}}{3+\frac{2}{x^2}+\frac{2}{x^3}}$$
Mit [[6]] folgern wir, dass
$$\lim\limits_{x \to \infty} \frac{1 + \frac{4}{x^2}}{3+\frac{2}{x^2}+\frac{2}{x^3}}= \frac{1}{3}$$
gilt. Eingesetzt in die Aufgabe erhalten wir:
$$\lim\limits_{x \to \infty} \ln\left(\frac{x^3+4x}{3x^3+2x^2+2x}\right)= \ln\left(\lim\limits_{x \to \infty}\frac{x^3+4x}{3x^3+2x^2+2x}\right)=\ln(\frac{1}{3})$$]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MA_LE09_22_Var</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>stetig</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>positiv</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>führende Terme</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>niedrigen</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>hohen</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>den Rechenregeln</text>
      <group>1</group>
    </dragbox>
  </question>

<!-- question: 1257778  -->
  <question type="multichoice">
    <name>
      <text>MAT_LE09_11_Stetigkeit_Def_umformulieren (Übung)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Welche der nachstehenden Gleichungen bedeuten, dass die Funktion \(f: \mathbb{R} \to \mathbb{R}\) im Punkt \(x=a\) stetig ist?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_11Ü</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(a)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(\lim\limits_{x \to a}x)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="0" format="html">
      <text><![CDATA[<p>\(f(a) = f(\lim\limits_{x \to a}x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="0" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a^+}f(x) =\lim\limits_{x \to a^-}f(x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257775  -->
  <question type="multichoice">
    <name>
      <text>Wahr/falsch 2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><span style="font-size: 0.9375rem;">Existiert der Grenzwert $\lim\limits_{x \to a}f(x)$ nicht, so auch nicht&nbsp;$\lim\limits_{x \to a^-}f(x)$ und&nbsp;$\lim\limits_{x \to a^+}f(x)$.</span><br></p><br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="0" format="html">
      <text><![CDATA[Wahr<br>]]></text>
      <feedback format="html">
        <text><![CDATA[<p>Die Aussage ist falsch, denn der rechts- und der linksseitige Grenzwert könnten existieren, müssen aber nicht gleich sein.&nbsp;</p>]]></text>
      </feedback>
    </answer>
    <answer fraction="100" format="html">
      <text><![CDATA[Falsch<br>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257774  -->
  <question type="multichoice">
    <name>
      <text>Wahr/falsch 3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Die Funktion $f(x) =&nbsp;\frac{1}{3-x}$ ist stetig auf ihrem Definitionsbereich \(D=\mathbb{R} \setminus \{3\}\).<br></p><br><p></p><p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="0" format="html">
      <text><![CDATA[<p>Falsch<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="100" format="html">
      <text><![CDATA[Wahr<br>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257779  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_19_Stetige Fortsetzbarkeit (Übung)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Entscheiden Sie für die folgenden Funktionen stetig fortsetzbar sind:<br>
<table>
<tr>
<td> Die Funktion \(f: \mathbb{R} \setminus\{2\} \to \mathbb{R}\) <br><br>
<img src="/theme/boost_union_vorsprung/pix/static/mathematik/le09/mat-le09-19b-grafik-beispiel-fuer-aufgabe-1.png" width="75%">
<br><br>
ist in \(x=2\) [[input:ans1]].[[validation:ans1]][[feedback:prt1]]</td>

<td> \(\quad\)    </td>

<td>Die Funktion \(f: \mathbb{R} \setminus\{0\} \to \mathbb{R}\) <br><br>
<img src="/theme/boost_union_vorsprung/pix/static/mathematik/le09/mat-le09-19c-grafik-beispiel-fuer-aufgabe-2.png" width="75%">
<br><br>
ist in \(x=0\) [[input:ans2]].[[validation:ans2]][[feedback:prt2]]</td>

</tr>
</table></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_19Ü</idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[L1: [[1,true, "stetig fortsetzbar"],[2,false,"nicht stetig fortsetzbar"]];
L2: [[1,false, "stetig fortsetzbar"],[2,true,"nicht stetig fortsetzbar"]];

L1:random_permutation(L1);
L2:random_permutation(L2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text>L1:random_permutation(L1);
L2:random_permutation(L2);
L3:random_permutation(L3);
L4:random_permutation(L4);
L5:random_permutation(L5);
L6:random_permutation(L6);</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>L1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>L2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Richtig, die Folge ist stetig fortsetzbar. Wir wählen in \(x=2\) als Funktionswert den Wert \(4\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Ihre Antwort ist leider falsch. </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Richtig, die Funktion ist nicht stetig fortsetzbar. Den Wert \(\infty\) können wir nicht als Funktionswert setzen.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Ihre Antwort ist leider falsch.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 1:</strong></p>

<p>Wiederholen Sie die Definition der stetigen Fortsetzbarkeit noch einmal.</p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 2:</strong></p>

<p>Welchen Punkt würden Sie im Graphen jeweils ergänzen wollen, um den Graphen durchzeichnen zu können?</p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p><strong>Hinweis 3 - Musterlösung:</strong></p>

<p> Die erste Funktion ist stetig fortsetzbar. In die Lücke bei \(x=2\) ergänzt man als Funktionswert den Wert \(4\).<br>
Die zweite Funktion ist nicht stetig fortsetzbar. Den Wert \(\infty\) können wir nicht als Funktionswert setzen.</p>]]></text>
    </hint>
  </question>

<!-- question: 839727  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_24_Parameterabhängige_Stetigkeit_rand</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>Bestimmen Sie \(a\) und \(b\) so, dass die Funktion</p>
<p>
\(f(x)&nbsp; = \left\{ \begin{array}{ll} {@fterm@} &amp; \textrm{ für }\; x\leq {@l@} \\
 {@gterm@} &amp;\textrm{ für }\; {@l@} &lt; x\leq {@m@} \\
{@hterm@} &amp; \textrm{ für }\;  x &gt; {@m@}
\end{array}\right.\)</p>
<p>stetig ist.
  <br>
</p>
<p>Antwort:
  <br>
</p>
<p>a = [[input:ans1]]
  <br>
</p>
<p>b = [[input:ans2]]
</p>
<p>[[validation:ans1]]
  <br>
</p>
<p>[[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text>j:1+rand(2);    /* bestimmt einen der Koeffizienten von f   */
k:1+rand(2);   /* bestimmt einen der Koeffizienten von h  */
l:-2+rand(2);   /* entscheidet, wo der Übergang von f zu g stattfindet   */
m:1+rand(2);   /* entscheidet, wo der Übergang von g zu h stattfindet   */
f(x):=x^2+j*a*x+b;
fterm:x^2+j*a*x+b;
 g(x):=(a+b)*x;
gterm:(a+b)*x;
 h(x):=x^2-a*x+k*b;
hterm:x^2-a*x+k*b;
c:flatten(solve([f(l)=g(l),g(m)=h(m)],[a,b]));
aloes:rhs(c[1]);
bloes:rhs(c[2]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote>
      <text>Stetigkeit</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Richtige Antwort, gut gemacht!</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Ihre Antwort ist teilweise korrekt.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>aloes</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>bloes</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>aloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(a={@aloes@}\;\textrm{ ist richtig bestimmt}\)<br></p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(a\;\textrm{ ist nicht richtig bestimmt}\)<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(b={@bloes@}\;\textrm{ ist richtig bestimmt}\)<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(b\;\textrm{ ist nicht richtig bestimmt}\)<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bloes</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>\(b={@bloes@}\;\textrm{ ist richtig bestimmt}\)<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(b\;\textrm{ ist nicht richtig bestimmt}\)<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>l^2+j*ans1*l+ans2-(ans1+ans2)*l</sans>
        <tans>0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@l@}\) stetig.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@l@}\) nicht stetig.<br></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>(ans1+ans2)*m-(m^2-ans1*m+k*ans2)</sans>
        <tans>0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@m@}\) stetig.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Funktion ist bei \(x={@m@}\) nicht stetig.<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE09/Variante nur im Test</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 1258152  -->
  <question type="multichoice">
    <name>
      <text>MAT_LE09_11_Stetigkeit_Def_umformulieren (Testvariante)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Welche der nachstehenden Gleichungen bedeuten, dass die Funktion \(f: \mathbb{R} \to \mathbb{R}\) im Punkt \(x=a\) stetig ist?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_11_Test</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text><![CDATA[Prima, Ihre Antwort ist richtig!<br>]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(a)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="50" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a}f(x) = f(\lim\limits_{x \to a}x)\)</p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="html">
      <text><![CDATA[<p>\(f(a) = f(\lim\limits_{x \to a} x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="-30" format="html">
      <text><![CDATA[<p>\(\lim\limits_{x \to a^+}f(x) =\lim\limits_{x \to a^-}f(x)\)<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <hint format="markdown">
      <text>**Hinweis 1:**

Schauen Sie sich noch einmal die Definition der Stetigkeit an und überprüfen, welche Unterschiede die Gleichungen zur Definition aufweisen.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 2:**

Überlegen Sie sich Beispiele für die einzelnen Gleichungen. Können Sie Funktionen konstruieren, die in $x=a$ nicht stetig sind?</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 3 - Musterlösung:**

Die Gleichungen
$$\lim\limits_{x \to a}f(x) = f(a)$$
 und
$$\lim\limits_{x \to a}f(x) = f(\lim\limits_{x \to a}x)$$
sind richtig.</text>
    </hint>
  </question>

<!-- question: 1257776  -->
  <question type="multichoice">
    <name>
      <text>Wahr/falsch 1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Sei $f: \mathbb{R} \to
\mathbb{R}$ eine Funktion. Gilt $\lim\limits_{x \to
a^-}f(x)=\lim\limits_{x \to a^+}f(x)$, so ist $f$ in $x=a$ stetig.</p><br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="0" format="html">
      <text><![CDATA[Wahr<br>]]></text>
      <feedback format="html">
        <text><![CDATA[<p>Die Aussage ist falsch, denn auch wenn der rechts- und der linksseitige Grenzwert gleich sind, könnte der Funktionswert $f(a)$ nicht mit diesen Grenzwerten übereinstimmen.</p>]]></text>
      </feedback>
    </answer>
    <answer fraction="100" format="html">
      <text><![CDATA[Falsch<br>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257773  -->
  <question type="multichoice">
    <name>
      <text>Wahr/falsch 5</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Sei $f: \mathbb{R} \to \mathbb{R}$ eine Funktion. Ist $f$ in $x=a$
stetig, so gilt $\lim\limits_{x \to a^-}f(x)=\lim\limits_{x \to
a^+}f(x)$.<br></p><p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MAT_LE09_Test</idnumber>
    <single>true</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="html">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="0" format="html">
      <text><![CDATA[<p>Falsch<br></p>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
    <answer fraction="100" format="html">
      <text><![CDATA[Wahr<br>]]></text>
      <feedback format="html">
        <text></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257780  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_24_Testvariante_Parameterabhängige_Stetigkeit</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>Bestimmen Sie \(A\), \(B\) und \(C\) so, dass die Funktion</p>
<p>
\(f(x)&nbsp; = \left\{ \begin{array}{ll} -4x-1 &amp; \textrm{ für }\; x\leq -1 \\
 Ax^2 +B &amp;\textrm{ für }\; -1 &lt; x\leq 1 \\
\ln(x)+C &amp; \textrm{ für }\;  x &gt; 1
\end{array}\right.\)</p>
<p>stetig ist.<br>
(Hinweis: Die Lösung ist nicht eindeutig.)
  <br>
</p>
<p>Antwort:
  <br>
</p>
<p>A = [[input:ans1]][[validation:ans1]]
  <br>
</p>
<p>B = [[input:ans2]][[validation:ans2]]
<br>
</p>
<p>C = [[input:ans3]][[validation:ans3]]
</p>
[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1+ans2</sans>
        <tans>3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sie haben \(A\) und \(B\) ist richtig bestimmt.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben die Werte für \(A\) und \(B\) leider nicht richtig bestimmt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sie haben \(C\) richtig bestimmt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben \(C\) leider nicht richtig bestimmt.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

<!-- question: 1257781  -->
  <question type="stack">
    <name>
      <text>MAT_LE09_Testfrage_Stetigkeitsstellen</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p>Gegeben sei die Funktion

\(f(x)&nbsp; = \left\{ \begin{array}{ll} \cos(x)+2 &amp; \textrm{ für }\; x\leq -\pi \\
 4(x+\pi)+1 &amp;\textrm{ für }\; -\pi &lt; x\leq \pi \\
|8x|-2 &amp; \textrm{ für }\;  x &gt; \pi
\end{array}\right.\)</p>

<p>Was können Sie über die Funktion im Bezug auf Stetigkeit aussagen? </p>
<p>Im Intervall \((-\infty, -\pi)\) ist die Funktion [[input:ans1]].[[validation:ans1]]</p>
<p>Im Intervall \((-\pi, \pi)\) ist die Funktion [[input:ans2]].[[validation:ans2]]</p>
<p>Im Intervall \((\pi, \infty)\) ist die Funktion [[input:ans3]].[[validation:ans3]]</p>
<p>Im Punkt \(x=-\pi\) ist die Funktion [[input:ans4]].[[validation:ans4]]</p>
<p>Im Punkt \(x=\pi\) ist die Funktion [[input:ans5]].[[validation:ans5]]</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[L1:[[1,true, "stetig"], [2,false, "nicht stetig"]];
L2:[[1,true, "stetig"], [2,false, "nicht stetig"]];
L3:[[1,true, "stetig"], [2,false, "nicht stetig"]];
L4:[[1,true, "stetig"], [2,false, "nicht stetig"]];
L5:[[1,false, "stetig"], [2,true, "nicht stetig"]];

L1:random_permutation(L1);
L2:random_permutation(L2);
L3:random_permutation(L3);
L4:random_permutation(L4);
L5:random_permutation(L5);]]></text>
    </questionvariables>
    <specificfeedback format="markdown">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote>
      <text>L1:random_permutation(L1);
L2:random_permutation(L2);
L3:random_permutation(L3);
L4:random_permutation(L4);
L5:random_permutation(L5);</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="markdown">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="markdown">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="markdown">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>L1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>L2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>L3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>L4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>L5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="markdown">
          <text>Die erste Antwort ist korrekt!</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="markdown">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="markdown">
          <text>Die zweite Antwort ist korrekt!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="markdown">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="markdown">
          <text>Die dritte Antwort ist korrekt!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="markdown">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="markdown">
          <text>Die vierte Antwort ist korrekt!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="markdown">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="markdown">
          <text>Die fünfte Antwort ist korrekt!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="markdown">
          <text>Nicht alle Ihre Antworten sind korrekt.</text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>