<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE06</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber>MA_LE06</idnumber>
  </question>

<!-- question: 1257396  -->
  <question type="ddwtos">
    <name>
      <text>MA_LE06_47_Verneinungen</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p> Ordnen Sie den Aussagen ihre <i>Negationen</i> korrekt zu:</p>

<p> \(A\): <i>Morgen werde ich aufräumen und einkaufen.</i> <br>Die Negation lautet \(\lnot A\): [[1]]</p>

<p>\(B\): <i>Am Wochenende werde ich einen Ausflug machen oder Freunde besuchen.</i> <br>Die Negation lautet \(\lnot B\): [[3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[Ihre Antwort ist richtig. <br>
Beim Negieren einer Aussage gilt \(\lnot (A \wedge B)=\lnot A \vee \lnot B\) und \(\lnot (A \vee B)=\lnot A \wedge \lnot B\).]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text><![CDATA[Die Antwort ist teilweise richtig.<br>
Beim Negieren einer Aussage gilt \(\lnot (A \wedge B)=\lnot A \vee \lnot B\) und \(\lnot (A \vee B)=\lnot A \wedge \lnot B\).]]></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text><![CDATA[Die Antwort ist falsch. <br> Beim Negieren einer Aussage gilt \(\lnot (A \wedge B)=\lnot A \vee \lnot B\) und \(\lnot (A \vee B)=\lnot A \wedge \lnot B\).]]></text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>Morgen werde ich nicht arbeiten und nicht einkaufen.</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>Morgen werde ich nicht arbeiten oder nicht einkaufen.</text>
      <group>1</group>
    </dragbox>
    <dragbox>
      <text>Am Wochenende werde ich keinen Ausflug machen und keine Freunde besuchen.</text>
      <group>2</group>
    </dragbox>
    <dragbox>
      <text>Am Wochenende werde ich keinen Ausflug machen oder keine Freunde besuchen.</text>
      <group>2</group>
    </dragbox>
  </question>

<!-- question: 1257403  -->
  <question type="ddwtos">
    <name>
      <text>MAT_LE06_41var_vollständige Induktion</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p> Beweisen Sie die folgende Aussage mit vollständiger Induktion: </p>
<p><center><i>\(n^3+2n\) ist durch \(3\) teilbar.</i></center></p><br>
<p> <strong>Beweis:</strong></p>
<p><i>Induktionsanfang:</i> Sei [[1]]. Dann gilt: [[2]]. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für \(n\) wahr: [[3]] ist durch \(3\) teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: [[4]] ist durch \(3\) teilbar.<br>
Wir formen diesen Term hierzu um:
<center>[[4]]\(=\)[[5]]\(+\)[[8]]\(=\)[[3]]\(+\)[[6]]</center>
Der erste Teilterm [[3]] ist nach Induktionsvoraussetzung durch \(3\) teilbar. Aus dem zweiten Teilterm können wir \(3\) ausklammern und erhalten: [[7]]. Damit ist der zweite Teilterm durch \(3\) teilbar.
Also ist der gesamte Term durch \(3\) teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_41var</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[<p> Sehr gut! Die Antwort ist richtig. </p>
[stashdrop secret="gO4hHY" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text><![CDATA[<p> Die Antwort ist teilweise richtig. Versuchen Sie es noch einmal! </p>]]></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text><![CDATA[<p> Ihre Antwort ist leider falsch. Versuchen Sie es noch einmal! </p>]]></text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>\(n=1\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 1^3+2\cdot 1=3 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\(n^3+2n\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( (n+1)^3+2(n+1) \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( n^3+3n^2+3n+1 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 3n^2+3n+3 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 3(n^2+n+1) \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 2n+2 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Schauen Sie noch einmal in die Beispiele. Was muss in der Induktionsvoraussetzung und dem Induktionsschritt gezeigt werden?</p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Setzen Sie $n=1$ in die Gleichung ein. Welches Ergebnis erhalten Sie? Ist diese Zahl durch $3$ teilbar?]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung </b> </p>
<p> Es gilt: </p>
<p><i>Induktionsanfang:</i> Sei $n=1$. Dann gilt: $1^3+2\cdot 1=3$. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für $n$ wahr: $n^3+2n$ ist durch 3 teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: $(n+1)^3+2(n+1)$ ist durch 3 teilbar.<br>
Wir formen diesen Term hierzu um:
<center>$(n+1)^3+2(n+1)=n^3+3n^2+3n+1 + 2n+2=(n^3+2n) + (3n^2+3n+3)</center>
Der erste Teilterm $n^3+2n$ ist nach Induktionsvoraussetzung durch $3$ teilbar. Aus dem zweiten Teilterm können wir $3$ ausklammern und erhalten: $3(n^2+n+1)$. Damit ist der zweite Teilterm durch $3$ teilbar.
Also ist der gesamte Term durch $3$ teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </hint>
  </question>

<!-- question: 1257392  -->
  <question type="ddwtos">
    <name>
      <text>MAT_LE06_41_vollständige Induktion</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p> Beweisen Sie die folgende Aussage mit vollständiger Induktion: </p>
<p><center><i>\(n^3+2n\) ist durch \(3\) teilbar.</i></center></p><br>
<p> <strong>Beweis:</strong></p>
<p><i>Induktionsanfang:</i> Sei [[1]]. Dann gilt: [[2]]. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für \(n\) wahr: [[3]] ist durch \(3\) teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: [[4]] ist durch \(3\) teilbar.<br>
Wir formen diesen Term hierzu um:
<center>[[4]]\(=\)[[5]]\(+\)[[8]]\(=\)[[3]]\(+\)[[6]]</center>
Der erste Teilterm [[3]] ist nach Induktionsvoraussetzung durch \(3\) teilbar. Aus dem zweiten Teilterm können wir \(3\) ausklammern und erhalten: [[7]]. Damit ist der zweite Teilterm durch \(3\) teilbar.
Also ist der gesamte Term durch \(3\) teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_41</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[<p> Sehr gut! Die Antwort ist richtig. </p>
[stashdrop secret="gO4hHY" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text><![CDATA[<p> Die Antwort ist teilweise richtig. Versuchen Sie es noch einmal! </p>]]></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text><![CDATA[<p> Ihre Antwort ist leider falsch. Versuchen Sie es noch einmal! </p>]]></text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>\(n=1\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 1^3+2\cdot 1=3 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\(n^3+2n\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( (n+1)^3+2(n+1) \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( n^3+3n^2+3n+1 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 3n^2+3n+3 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 3(n^2+n+1) \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 2n+2 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Schauen Sie noch einmal in die Beispiele. Was muss in der Induktionsvoraussetzung und dem Induktionsschritt gezeigt werden?</p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Setzen Sie $n=1$ in die Gleichung ein. Welches Ergebnis erhalten Sie? Ist diese Zahl durch $3$ teilbar?]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung </b> </p>
<p> Es gilt: </p>
<p><i>Induktionsanfang:</i> Sei $n=1$. Dann gilt: $1^3+2\cdot 1=3$. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für $n$ wahr: $n^3+2n$ ist durch 3 teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: $(n+1)^3+2(n+1)$ ist durch 3 teilbar.<br>
Wir formen diesen Term hierzu um:
<center>$(n+1)^3+2(n+1)=n^3+3n^2+3n+1 + 2n+2=(n^3+2n) + (3n^2+3n+3)</center>
Der erste Teilterm $n^3+2n$ ist nach Induktionsvoraussetzung durch $3$ teilbar. Aus dem zweiten Teilterm können wir $3$ ausklammern und erhalten: $3(n^2+n+1)$. Damit ist der zweite Teilterm durch $3$ teilbar.
Also ist der gesamte Term durch $3$ teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </hint>
  </question>

<!-- question: 1257393  -->
  <question type="gapselect">
    <name>
      <text>Äquivalenz von Aussagen</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Verifizieren Sie anhand der folgenden Wahrheitstafel, dass die beiden Aussagen \((A\Rightarrow B)\Rightarrow B\) und \(A\vee B\) äquivalent zueinander sind.
</p>
<table style="border: 3px solid; border-spacing:10px;">

  <tbody>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\; \;\;\; A\;\;\; \;\)</td>
      <td style="border: 2px solid">\(\; \;\; \;B\;\; \;\;\)</td>
      <td style="border: 2px solid">\(\; \; A\Rightarrow B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; (A\Rightarrow B)\Rightarrow B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; A\vee B\; \;\)</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
      <td style="border: 2px solid">\(\;\)[[8]]</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
    </tr>

  </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>wahr</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>4</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>4</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257439  -->
  <question type="gapselect">
    <name>
      <text>MAT_LE06_12_Negation</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p>Vergleichen Sie jetzt Ihre Lösungen mit diesen Möglichkeiten zur Auswahl: <br>
<ul>
<li>Aussage \(A\): \(2>7\). <br>Die Negation ist \(\lnot A\): [[1]]</li>

<li>Aussage \(B\): Die Summe aus 3 und 5 ist 8. <br>Die Negation ist \(\lnot B\): [[4]] </li>
<li>Aussage \(C\): Es gibt eine natürliche Zahl zwischen \(15\) und \(20\), die durch \(7\) teilbar ist. <br>Die Negation ist \(\lnot C\): [[7]] </li>
</ul></p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_12</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[Die Antwort ist richtig!
[stashdrop secret="zXQ59p" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>\(2 \leq 7\)</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text><![CDATA[\(2 < 7\)]]></text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text><![CDATA[\(2&lt;4\) und \(4&lt;7\)]]></text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>Die Summe aus \(3\) und \(5\) ist eine andere Zahl als \(8\).</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>Die Differenz aus \(5\) und \(2\) ist eine andere Zahl als \(8\).</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>Die Summe aus \(3\) und \(5\) ist \(9\).</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>Alle natürlichen Zahlen zwischen \(15\) und \(20\) sind nicht durch \(7\) teilbar.</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>Eine natürliche Zahl zwischen \(15\) und \(20\) ist nicht durch \(7\) teilbar.</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>Es gibt eine natürliche Zahl zwischen \(15\) und \(20\), die durch \(8\) teilbar ist.</text>
      <group>3</group>
    </selectoption>
  </question>

<!-- question: 1257347  -->
  <question type="multichoice">
    <name>
      <text>MA_LE06_05_Beispiele_Aussagen</text>
    </name>
    <questiontext format="markdown">
      <text>Welche der folgenden Sätze sind Aussagen?</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MA_LE06_05</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="markdown">
      <text><![CDATA[Die Antwort ist richtig.

[stashdrop secret="3ubl3R" image]]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="50" format="markdown">
      <text>20 ist durch 5 teilbar.</text>
      <feedback format="markdown">
        <text>Sehr gut, dies ist eine Aussage. Es ist zudem eine wahre Aussage.</text>
      </feedback>
    </answer>
    <answer fraction="-50" format="markdown">
      <text>Komm, lasst uns zusammen eine Pause machen!</text>
      <feedback format="markdown">
        <text>Dieser Satz ist keine Aussage, es ist ein Ausruf.</text>
      </feedback>
    </answer>
    <answer fraction="50" format="markdown">
      <text>Der Februar hat immer 28 Tage.</text>
      <feedback format="markdown">
        <text>Dieser Satz ist eine Aussage. Es ist zudem eine falsche Aussage.</text>
      </feedback>
    </answer>
    <hint format="markdown">
      <text>**Hinweis 1:**

Lesen Sie sich die Definition einer Aussage nochmals durch.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 2:**

Schauen Sie sich das Beispiel vor dieser Aufgabe nochmals an.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 3:**

Der Satz *$20$ ist durch $5$ teilbar.*  ist eine wahre Aussage. 
 
Der Satz *Der Februar hat immer $28$ Tage.* ist eine falsche Aussage. 

Der Satz *Komm, lasst uns zusammen eine Pause machen!* ist ein Ausruf und keine Aussage.</text>
    </hint>
  </question>

<!-- question: 1257438  -->
  <question type="shortanswer">
    <name>
      <text>LE06_36L_Aufloesung</text>
    </name>
    <questiontext format="markdown">
      <text>Fügen Sie hier das Lösungswort ein, um eine Münze zu erhalten:</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_36L</idnumber>
    <usecase>0</usecase>
    <answer fraction="100" format="moodle_auto_format">
      <text>Beweis</text>
      <feedback format="markdown">
        <text><![CDATA[Richtig!
[stashdrop secret="crerib" image]

<p>Wir beweisen die Aussage:
<center><i>\(x\) durch \(5\) teilbar \(\Rightarrow\) \(x^2\) ist durch \(5\) teilbar.</i> </center><br>

<strong>Beweis:</strong><br>
Sei \(x\) durch \(5\) teilbar. Dann lässt sich \(x\) als \(x=5\cdot r\) mit einer natürlichen Zahl \(r\) darstellen. Wir quadrieren \(x\). Dann erhalten wir \(x^2=(5\cdot r)^2=5\cdot r \cdot 5\cdot r= 5(5\cdot r^2)\). \(r^2\) und \(5\cdot r^2\) sind weiterhin natürliche Zahlen. Damit ist \(x^2\) durch \(5\) teilbar. \(\Box\)</p>]]></text>
      </feedback>
    </answer>
  </question>

<!-- question: 1257348  -->
  <question type="stack">
    <name>
      <text>MAT_LE06_24_Aussageformen</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Aussageformen können von mehreren Variablen abhängen. Wir betrachten die
Aussageform \(P(a,b,c):\) \(a^2+b^2=c^2\) mit \(a,b,c \in \mathbb{N}\).<br>
Bestimmen Sie den Wahrheitswert der folgenden Aussagen:</p>

<p>Die Aussage \(P(3,4,5)\) ist  [[input:ans1]] [[validation:ans1]] </p>
<p>Die Aussage \(P(1,1,1)\) ist  [[input:ans2]] [[validation:ans2]] </p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_24</idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[L1: [[1, true,"wahr"],[2, false,"falsch"]];
L2: [[1, false,"wahr"],[2, true,"falsch"]];
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>L1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>+</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>L2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>Aussage_1_W</trueanswernote>
        <truefeedback format="markdown">
          <text><![CDATA[<p>Sehr gut! Die Antwort ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Aussage_1_F</falseanswernote>
        <falsefeedback format="markdown">
          <text><![CDATA[<p> Ihre Antwort ist leider falsch. Die Aussage ist wahr. </p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Aussage_2_W</trueanswernote>
        <truefeedback format="markdown">
          <text><![CDATA[<p>Sehr gut! Die Antwort ist richtig.</p>

[stashdrop secret="pJ5nIi" image]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Aussage_2_F</falseanswernote>
        <falsefeedback format="markdown">
          <text><![CDATA[<p> Ihre Antwort ist leider falsch. Die Aussage \(P(3,4,5)\) ist wahr. </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(L1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(L2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_W</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_W</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_F</expectedanswernote>
      </expected>
    </qtest>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Setzen Sie die Zahlenwerte in die Definition der Aussageform \(P\) ein. </p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Sie setzen die Zahlenwerte in die Aussageform \(P\) ein und erhalten für die erste Frage: \(3^2+4^2=5^2\). Diese Aussage ist wahr. Probieren Sie es nun mit der zweiten Frage analog aus.</p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung: </b> </p>
<p>
Die erste Aussage \(P(3,4,5)\) ist wahr, denn \(3^2+4^2+5^2\) ist wahr<br>
Die zweite Aussage \(P(1,1,1)\) ist falsch, denn \(1^2+1^2 \not= 1^2\)
</p>]]></text>
    </hint>
  </question>

<!-- question: 1257391  -->
  <question type="stack">
    <name>
      <text>MAT_LE06_8_Konjunktion</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Die Aussage \(A \wedge B\) mit den Aussagen \(A\): <i>\(8\) ist eine ungerade
Zahl</i> und \(B\): <i>Berlin liegt in Australien</p> ist [[input:ans1]]. [[validation:ans1]]</p>

<p> Die Aussage \(C \wedge D\) mit den Aussagen \(C\): <i>\(25\) ist eine Quadratzahl</i> und \(D\): <i>\(3+5=8\)</p> ist [[input:ans2]]. [[validation:ans2]]</p>

<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_8</idnumber>
    <stackversion>
      <text>2023010400</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[L2: [[1, true,"wahr"],[2, false,"falsch"]];
L1: [[1, false,"wahr"],[2, true,"falsch"]];
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt2]]</p>]]></text>
    </specificfeedback>
    <questionnote>
      <text></text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>(</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>L1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>+</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>L2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>Aussage_1_W</trueanswernote>
        <truefeedback format="markdown">
          <text><![CDATA[<p>Sehr gut! Die erste Antwort ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>Aussage_1_F</falseanswernote>
        <falsefeedback format="markdown">
          <text><![CDATA[<p> Ihre Antwort ist leider falsch. Da beide Aussagen falsch sind, ist auch \(A \wedge B\) falsch. </p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Aussage_2_W</trueanswernote>
        <truefeedback format="markdown">
          <text><![CDATA[<p>Sehr gut! Die zweite Antwort ist richtig.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Aussage_2_F</falseanswernote>
        <falsefeedback format="markdown">
          <text><![CDATA[<p> Ihre Antwort ist leider falsch. Da beide Aussagen wahr sind, ist auch \(C \wedge D\) wahr.  </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Prima!</p>
[stashdrop secret="Vx2ViI" image]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(L1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(L2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_W</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_W</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Aussage_2_F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Schauen Sie sich noch einmal die Wahrheitswerttabelle für die Konjunktion an. Welche Zeile trifft auf die beiden Aussagen \(A\) und \(B\) zu? </p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Für die erste Aussagen nutzen wir, dass wir bereits wissen, dass \(A\) und \(B\) wahre Aussagen sind.</p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung: </b> </p>
<p>
Die erste Aussage \(A \wedge B\) ist wahr, denn sowohl \(A\) als auch \(B\) sind wahre Aussagen.<br>
Die zweite Aussage \(C \wedge D\) ist falsch, denn \(C\) ist eine wahre, aber \(D\) ist eine falsche Aussage.
</p>]]></text>
    </hint>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE06/Aufgaben zur Übung</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 1257394  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_45_Oder-Verknüpfung</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Unter einer <i>Disjunktion</i> von zwei Aussagen $A$ und $B$ verstehen wir die Aussage $A \vee B$ (wir sagen: <i>$A$ oder $B$</i>), die genau dann wahr ist, wenn wenigstens eine der beiden Aussagen wahr ist.</p>
<p>Vervollständigen Sie die Wahrheitswerttabelle dieser Verknüpfung:</p><br>
<table style="border: 3px solid; border-spacing:10px;">

  <tbody>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\; \;\;\; A\;\;\; \;\)</td>
      <td style="border: 2px solid">\(\; \;\; \;B\;\; \;\;\)</td>
      <td style="border: 2px solid">\(\; \; A\vee B\; \;\)</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[4]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[3]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
    </tr>

  </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_45</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>1</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257395  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_46_Quantoren</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Bestimmen Sie den Wahrheitswert der folgenden Aussagen:</p><br>

<p> <i>Für alle \(x \in \mathbb{R}\) gilt:  \(x^2 = 9\).</i> Diese Aussage ist [[1]].</p>

<p> <i>Es existiert ein \(x \in \mathbb{R}\) mit:  \(x^2 = 9\).</i> Diese Aussage ist [[2]].</p>

<p> <i>Es existiert genau ein \(x \in \mathbb{R}\) mit:  \(x^2 = 9\).</i> Diese Aussage ist [[3]].</p>

<p> <i>Es existiert genau ein \(x \in \mathbb{N}\) mit: \(x^2 = 9\).</i> Diese Aussage ist [[4]].</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_46</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>1</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257397  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_47_Äquivalenz</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Weisen Sie über die Wahrheitswerttabelle nach, dass die die Aussagen $\lnot(A \wedge B)$ und $\lnot A \vee \lnot B$ äquivalent sind: </p><br>
<table style="border: 3px solid; border-spacing:10px;">

  <tbody>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\; \;\;\; A\;\;\; \;\)</td>
      <td style="border: 2px solid">\(\; \;\; \;B\;\; \;\;\)</td>
      <td style="border: 2px solid">\(\; \; A\wedge B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot(A\wedge B)\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot A\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot A \vee \lnot B\; \;\)</td>

    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
      <td style="border: 2px solid">\(\;\)[[3]]</td>
      <td style="border: 2px solid">\(\;\)[[4]]</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[6]]</td>
      <td style="border: 2px solid">\(\;\)[[7]]</td>
      <td style="border: 2px solid">\(\;\)[[8]]</td>
      <td style="border: 2px solid">\(\;\)[[9]]</td>
      <td style="border: 2px solid">\(\;\)[[10]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[11]]</td>
      <td style="border: 2px solid">\(\;\)[[12]]</td>
      <td style="border: 2px solid">\(\;\)[[13]]</td>
      <td style="border: 2px solid">\(\;\)[[14]]</td>
      <td style="border: 2px solid">\(\;\)[[15]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[16]]</td>
      <td style="border: 2px solid">\(\;\)[[17]]</td>
      <td style="border: 2px solid">\(\;\)[[18]]</td>
      <td style="border: 2px solid">\(\;\)[[19]]</td>
      <td style="border: 2px solid">\(\;\)[[20]]</td>
    </tr>

  </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_47</idnumber>
    <shuffleanswers>0</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>4</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>5</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>11</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>13</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>4</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>5</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>8</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>11</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>6</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>9</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>14</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>10</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>8</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>9</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>10</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>13</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>20</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>19</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 0  -->
  <question type="category">
    <category>
      <text>$course$/top/MA_LE06/Variante nur im Test</text>
    </category>
    <info format="markdown">
      <text></text>
    </info>
    <idnumber></idnumber>
  </question>

<!-- question: 1257402  -->
  <question type="ddwtos">
    <name>
      <text>MAT_LE06_41b_vollständige Induktion (Testvariante)</text>
    </name>
    <questiontext format="markdown">
      <text><![CDATA[<p> Beweisen Sie die folgende Aussage mit vollständiger Induktion: </p>
<p><center><i>\(5^n+7\) ist durch \(4\) teilbar.</i></center></p><br>
<p> <strong>Beweis:</strong></p>
<p><i>Induktionsanfang:</i> Sei [[1]]. Dann gilt: [[2]]. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für \(n\) wahr: [[3]] ist durch \(4\) teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: [[4]] ist durch \(4\) teilbar.<br>
Wir formen diesen Term hierzu um:
<center>[[4]]\(=\)[[5]]\(=\)[[6]]\(=\)[[7]]</center>
Der zweite Teilterm [[3]] ist nach Induktionsvoraussetzung durch \(4\) teilbar. Der erste Teilterm enthält den Faktor \(4\). Daher ist er durch \(4\) teilbar.
Also ist der gesamte Term durch \(4\) teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>LE06_42</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="markdown">
      <text><![CDATA[<p> Sehr gut! Die Antwort ist richtig. </p>]]></text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text><![CDATA[<p> Die Antwort ist teilweise richtig. Versuchen Sie es noch einmal! </p>]]></text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text><![CDATA[<p> Ihre Antwort ist leider falsch. Versuchen Sie es noch einmal! </p>]]></text>
    </incorrectfeedback>
    <shownumcorrect/>
    <dragbox>
      <text>\(n=1\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 5^1+7=12 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\(5^n+7\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\(5^{n+1}+7\)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 5^n\cdot 5 + 7 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 5^n\cdot (4+1) + 7  \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <dragbox>
      <text>\( 4 \cdot 5^n + 5^n + 7 \)</text>
      <group>1</group>
      <infinite/>
    </dragbox>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 1: </b> </p>
<p> Schauen Sie noch einmal in die Beispiele. Was muss in der Induktionsvoraussetzung und dem Induktionsschritt gezeigt werden?</p>]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 2: </b> </p>
<p> Setzen Sie $n=1$ in die Gleichung ein. Welches Ergebnis erhalten Sie? Ist diese Zahl durch $3$ teilbar?]]></text>
    </hint>
    <hint format="markdown">
      <text><![CDATA[<p> <b> Hinweis 3 - Musterlösung </b> </p>
<p> Es gilt: </p>
<p><i>Induktionsanfang:</i> Sei $n=1$. Dann gilt: $1^3+2\cdot 1=3$. Dann ist die Induktionsvoraussetzung erfüllt.<br>
<i>Induktionsvoraussetzung:</i> Aussage ist für $n$ wahr: $n^3+2n$ ist durch 3 teilbar.<br>
<i>Induktionsschritt:</i> Wir zeigen: $(n+1)^3+2(n+1)$ ist durch 3 teilbar.<br>
Wir formen diesen Term hierzu um:
<center>$(n+1)^3+2(n+1)=n^3+3n^2+3n+1 + 2n+2=(n^3+2n) + (3n^2+3n+3)</center>
Der erste Teilterm $n^3+2n$ ist nach Induktionsvoraussetzung durch $3$ teilbar. Aus dem zweiten Teilterm können wir $3$ ausklammern und erhalten: $3(n^2+n+1)$. Damit ist der zweite Teilterm durch $3$ teilbar.
Also ist der gesamte Term durch $3$ teilbar.<br>
Damit ist die Aussage gezeigt.</p>]]></text>
    </hint>
  </question>

<!-- question: 1257398  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_46b_Quantoren (Testfrage)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Bestimmen Sie den Wahrheitswert der folgenden Aussagen:</p><br>

<p> <i>Für alle \(x \in \mathbb{R}\) gilt:  \(x^2 = 0\).</i> Diese Aussage ist [[1]].</p>

<p> <i>Es existiert ein \(x \in \mathbb{R}\) mit:  \(x^2 = 0\).</i> Diese Aussage ist [[2]].</p>

<p> <i>Es existiert genau ein \(x \in \mathbb{R}\) mit:  \(x^2 = 0\).</i> Diese Aussage ist [[3]].</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_48</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>1</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257400  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_47_Äquivalenz (Testvariante)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Weisen Sie über die Wahrheitswerttabelle nach, dass die die Aussagen $\lnot(A \vee B)$ und $\lnot A \wedge \lnot B$ äquivalent sind: </p><br>
<table style="border: 3px solid; border-spacing:10px;">

  <tbody>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\; \;\;\; A\;\;\; \;\)</td>
      <td style="border: 2px solid">\(\; \;\; \;B\;\; \;\;\)</td>
      <td style="border: 2px solid">\(\; \; A\vee B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot(A\vee B)\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot A\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot B\; \;\)</td>
      <td style="border: 2px solid">\(\; \; \lnot A \wedge \lnot B\; \;\)</td>

    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
      <td style="border: 2px solid">\(\;\)[[3]]</td>
      <td style="border: 2px solid">\(\;\)[[4]]</td>
      <td style="border: 2px solid">\(\;\)[[5]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[6]]</td>
      <td style="border: 2px solid">\(\;\)[[7]]</td>
      <td style="border: 2px solid">\(\;\)[[8]]</td>
      <td style="border: 2px solid">\(\;\)[[9]]</td>
      <td style="border: 2px solid">\(\;\)[[10]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[11]]</td>
      <td style="border: 2px solid">\(\;\)[[12]]</td>
      <td style="border: 2px solid">\(\;\)[[13]]</td>
      <td style="border: 2px solid">\(\;\)[[14]]</td>
      <td style="border: 2px solid">\(\;\)[[15]]</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[16]]</td>
      <td style="border: 2px solid">\(\;\)[[17]]</td>
      <td style="border: 2px solid">\(\;\)[[18]]</td>
      <td style="border: 2px solid">\(\;\)[[19]]</td>
      <td style="border: 2px solid">\(\;\)[[20]]</td>
    </tr>

  </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_48_7</idnumber>
    <shuffleanswers>0</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>4</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>5</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>6</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>7</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>4</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>8</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>5</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>9</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>6</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>10</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>11</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>12</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>7</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>8</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>9</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>10</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>11</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>12</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257399  -->
  <question type="gapselect">
    <name>
      <text>MA_LE06_48_Ergänzung (Test)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Vervollständigen Sie die Wahrheitswerttabelle dieser Verknüpfung:</p><br>
<table style="border: 3px solid; border-spacing:10px;">

  <tbody>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\; \;\;\; A\;\;\; \;\)</td>
      <td style="border: 2px solid">\(\; \;\; \;B\;\; \;\;\)</td>
      <td style="border: 2px solid">\(\; \; A\wedge B\; \;\)</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[1]]\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[4]]\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)[[3]]\(\;\)</td>
      <td style="border: 2px solid">\(\;\)wahr\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch</td>
    </tr>
    <tr style="border: 2px solid">
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)falsch\(\;\)</td>
      <td style="border: 2px solid">\(\;\)[[2]]</td>
    </tr>

  </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>3.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber>LE_06_48_1</idnumber>
    <shuffleanswers>1</shuffleanswers>
    <correctfeedback format="html">
      <text>Alle Antworten sind richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="html">
      <text>Ihre Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="html">
      <text>Alle Antworten sind falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <selectoption>
      <text>wahr</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>2</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>1</group>
    </selectoption>
    <selectoption>
      <text>falsch</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>3</group>
    </selectoption>
    <selectoption>
      <text>wahr</text>
      <group>1</group>
    </selectoption>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
    <hint format="html">
      <text></text>
      <shownumcorrect/>
    </hint>
  </question>

<!-- question: 1257401  -->
  <question type="multichoice">
    <name>
      <text>MA_LE06_05b_Beispiele_Aussagen (Test)</text>
    </name>
    <questiontext format="markdown">
      <text>Welche der folgenden Sätze sind Aussagen?</text>
    </questiontext>
    <generalfeedback format="markdown">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber>MA_LE06_06</idnumber>
    <single>false</single>
    <shuffleanswers>true</shuffleanswers>
    <answernumbering>none</answernumbering>
    <showstandardinstruction>1</showstandardinstruction>
    <correctfeedback format="markdown">
      <text>Die Antwort ist richtig.</text>
    </correctfeedback>
    <partiallycorrectfeedback format="markdown">
      <text>Die Antwort ist teilweise richtig.</text>
    </partiallycorrectfeedback>
    <incorrectfeedback format="markdown">
      <text>Die Antwort ist falsch.</text>
    </incorrectfeedback>
    <shownumcorrect/>
    <answer fraction="50" format="markdown">
      <text>21 ist keine Primzahl.</text>
      <feedback format="markdown">
        <text>Sehr gut, dies ist eine Aussage. Es ist zudem eine wahre Aussage.</text>
      </feedback>
    </answer>
    <answer fraction="-50" format="markdown">
      <text>Sollen wir zusammen etwas kochen?</text>
      <feedback format="markdown">
        <text>Dieser Satz ist keine Aussage, es ist eine Frage.</text>
      </feedback>
    </answer>
    <answer fraction="50" format="markdown">
      <text>Hamburg ist die Hauptstadt von Deutschland.</text>
      <feedback format="markdown">
        <text>Dieser Satz ist eine Aussage. Es ist zudem eine falsche Aussage.</text>
      </feedback>
    </answer>
    <hint format="markdown">
      <text>**Hinweis 1:**

Lesen Sie sich die Definition einer Aussage nochmals durch.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 2:**

Schauen Sie sich das Beispiel vor dieser Aufgabe nochmals an.</text>
    </hint>
    <hint format="markdown">
      <text>**Hinweis 3:**

Der Satz *$20$ ist durch $5$ teilbar.*  ist eine wahre Aussage. 
 
Der Satz *Der Februar hat immer $28$ Tage.* ist eine falsche Aussage. 

Der Satz *Komm, lasst uns zusammen eine Pause machen!* ist ein Ausruf und keine Aussage.</text>
    </hint>
  </question>

</quiz>